## Exercise 03.1

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try check for equality, e.g. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
# checking for round of error for decimals
a = 100
b = 0.1
c = 0.2
d0 = a*b + a*c
d1 = a*(b + c)

print("%.30f"%(d0 - d1))
print(d0 == d1)
```

Output:
```
-0.000000000000003552713678800501
False
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

## Exercise 03.2

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (a) the compact form $(x + y)^{6}$; and (b) the expanded form for:

- $x = 10$ and $y = 10.1$
- $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the first case using the variable `f0` **this should be _f1_ I bet?**, and using the variable `f1` **_f2_?**for the second case.

For the second case, compare the computed and analytical solutions.
Which approach would you recommend for computing this expression?

#### (a) $x = 10$ and $y = 10.1$

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
def compact_f(x, y):
    return (x + y)**6

def expanded_f(x, y):
    return x**6 + 6*x**5*y + 15*x**4*y**2 + 20*x**3*y**3 + 15*x**2*y**4 + 6*x*y**5 + y**6

x = 10.0
y = 10.1

# testing the results
f1 = compact_f(x, y)
f0 = expanded_f(x, y)
print("f0 is {0:.30f} and f1 is {1:.30f}".format(f0, f1))
```

Output:
```
f0 is 65944160.601200997829437255859375000000 and f1 is 65944160.601201027631759643554687500000
```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
assert round(f0 - 65944160.601201, 10) == 0.0
assert round(f1 - 65944160.60120103, 10) == 0.0
```

#### (b) $x = 10$ and $y = -10.1$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

# testing the results
f1 = compact_f(x, y)
f0 = expanded_f(x, y)
print("f0 is {0:.30f} and f1 is {1:.30f}".format(f0, f1))
```

Output:
```
f0 is 0.000000995816662907600402832031 and f1 is 0.000000999999999999978778924430
```

```python
assert round(f0 - 1.0e-6, 2) == 0.0
assert round(f1 - 1.0e-6, 8) == 0.0
```

## Exercise 03.3

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

- $x = 1 \times 10^{7}$
- $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
import numpy as np

def f_origin(x):
    return 1 / (np.sqrt(x**2 - 1) - x)

def f_altered(x):
    return - np.sqrt(x**2 - 1) - x

def display_both(x):
    print("x is: ", x, sep = " ")
    print("origin function gives: {0:.30f} \naltered function gives: {1:.30f}".format(f_origin(x), f_altered(x)))
    
display_both(1e7)
```

Output:
```
x is:  10000000.0
origin function gives: -19884107.851851850748062133789062500000 
altered function gives: -19999999.999999947845935821533203125000
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
display_both(1e8)
```

Output:
```
x is:  100000000.0
origin function gives: inf 
altered function gives: -200000000.000000000000000000000000000000
/home/nbuser/anaconda3_410/lib/python3.5/site-packages/ipykernel/__main__.py:4: RuntimeWarning: divide by zero encountered in double_scalars
```

```python

```