## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    # YOUR CODE HERE
    if n < 0:
        raise ValueError("index to the function may not be negative")
    elif n == 0:
        return 0
    elif n == 1:
        return 1
    else:
        return f(n-1) + f(n-2)
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12

def compute_root(f, x0, x1, tol, max_it):
    "Compute roots of a function using bisection"
    if max_it > 10:
        raise ValueError("TooManyIterationRequiredByTheUserWhoIsTooInconciderate: the free plan you are in only support 10 iteration..")
    x_mid = (x0 + x1) / 2
    f_mid = f(x_mid)
    if abs(f_mid) < tol:
        it = 1
    elif f_mid * f(x0) <= 0:
        x_mid, f_mid, it = compute_root(f, x0, x_mid, tol, max_it)
        it += 1
    else:
        x_mid, f_mid, it = compute_root(f, x_mid, x1, tol, max_it)
        it += 1
    
    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=20)
```

Output:
```
<ipython-input-10-ac365e77e568> in <module>()
      1 # Test with max_it = 30
----> 2 x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)
      3 
      4 # Test with max_it = 20
      5 with pytest.raises(RuntimeError):

<ipython-input-9-4d2f0fea7548> in compute_root(f, x0, x1, tol, max_it)
      6     "Compute roots of a function using bisection"
      7     if max_it > 10:
----> 8         raise ValueError("TooManyIterationRequiredByTheUserWhoIsTooInconciderate: the free plan you are in only support 10 iteration..")
      9     x_mid = (x0 + x1) / 2
     10     f_mid = f(x_mid)

ValueError: TooManyIterationRequiredByTheUserWhoIsTooInconciderate: the free plan you are in only support 10 iteration..
```

```python

```